```python
# Initialize Otter
import otter
grader = otter.Notebook("nearest_neighbors.ipynb")
```

# **Домашнее задание 2.** Метод K-ближайших соседей

<p style="color:red">  
<b>Предупреждение:</b> в данном задании разрешено использовать только NumPy, Pandas и библиотеки для визуализации! 
</p>

На семинаре мы познакомились с реализацией KNN в библиотеке scikit-learn, обучили модели для задач классификации и регрессии и оценили их качество на синтетических данных.

Следуя высказыванию одного удивительного человека:

> «What I cannot create, I do not understand.» — Richard Feynman

в этом задании мы реализуем алгоритм k-ближайших соседей на чистом NumPy. Приступим!

```python
from typing import Any, Literal, Self

import numpy as np
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
```

```python
def convert_pandas_to_numpy(func: callable) -> callable:
    def wrapper(*args: tuple[Any], **kwargs: dict[Any]):
        args = tuple(arg.values if isinstance(arg, (pd.DataFrame, pd.Series)) else arg for arg in args)
        kwargs = {
            key: arg.values if isinstance(arg, (pd.DataFrame, pd.Series)) else arg
            for key, arg in kwargs.items()
        }  # fmt: skip

        return func(*args, **kwargs)

    return wrapper
```

### **Задание 0.** Предобработка данных

Прочитайте данные из файла `Predict_Calorie_Expenditure_Dataset.csv`, при необходимости удалите пропущенные значения, проведите разведочный анализ данных.  

_Данное задание не оценивается, но его выполнение необходимо для получения полного балла._

_Points:_ 0

```python
data = ...
data = data.sample(5_000, random_state=42)

print(f"{data.shape = }")
data.head()
```

```python
grader.check("Task0")
```

Разобьем данные на обучающую, валидационную и тестовые подвыборки

```python
X, y = data.drop(columns="Calories"), data["Calories"]

X_test = X.sample(frac=0.2, random_state=42)
test_idx = X_test.index
y_test = y.loc[test_idx]

train_idx = X.index.difference(test_idx)
X_train = X.loc[train_idx]
y_train = y.loc[train_idx]

print(f"Размер обучающей выборки:\t{y_train.size} объектов")
print(f"Размер тестовой выборки:\t{y_test.size} объектов")
```

### **Задание 1.** Евклидово расстояние

<p style="color:red">  
<b>Предупреждение:</b> использование циклов <code>for</code> и <code>while</code> в данном задании запрещено!  
</p>  

Пусть матрица $ X_A \in \mathbb{R}^{n \times d} $ содержит объекты с известными целевыми переменными, а матрица $ X_B \in \mathbb{R}^{m \times d} $ — объекты, для которых мы хотим сделать предсказания. Для применения модели KNN нам необходимо научиться считать матрицу расстояний $ D \in \mathbb{R}^{n \times m} $.

В теории мы могли бы, используя вложенные циклы, пройти по строкам первой и второй матриц, считая расстояния между ними, но ярко-красное предупреждение выше не рекомендует нам этого делать. Значит, будем искать другой способ, а для этого посмотрим внимательнее на определение евклидова расстояния:

$$
d(\mathbf{v}, \mathbf{u}) = \sqrt{\sum_{i=1}^{d} (v_i - u_i)^2}, \quad \mathbf{v},\, \mathbf{u} \in \mathbb{R}^d.
$$

Теперь избавимся от корня:

$$
\begin{aligned}
d^2(\mathbf{v}, \mathbf{u})
&= \sum_{i=1}^{d} (v_i - u_i)^2 \\
&= \sum_{i=1}^{d} v_i^2 + \sum_{i=1}^{d} u_i^2 - 2 \sum_{i=1}^{d} v_i u_i \\
&= \lVert \mathbf{v} \rVert^2 + \lVert \mathbf{u} \rVert^2 - 2\,\mathbf{v}\cdot\mathbf{u}.
\end{aligned}
$$

Иными словами, для расчета евклидова расстояния в матричной форме нам всего-то и нужно, что посчитать нормы векторов в исходных матрицах и найти матрицу скалярных произведений. С этим уже можно работать!

**Подсказка:** рассмотрите произведение $ X_A X_B^\top $.

Реализуйте функцию `cdist`, которая вычисляет евклидовы расстояния между строками двух матриц.

_Points:_ 5

```python
@convert_pandas_to_numpy
def cdist(XA: np.ndarray, XB: np.ndarray) -> np.ndarray:
    """Вычисляет матрицу евклидовых расстояний между векторами из двух входных матриц.

    Args:
        XA (np.ndarray): Матрица размера (n, d), содержащая n d-мерных векторов.
        XB (np.ndarray): Матрица размера (m, d), содержащая m d-мерных векторов.

    Returns:
        np.ndarray: Матрица расстояний размера (n, m).
    """
    ...
```

```python
grader.check("Task1")
```

### **Задание 2.** Базовый класс

Наконец, мы готовы приступить к имплементации KNN. Сначала создадим базовый класс, который будет содержать логику обучения и поиска ближайших соседей.

_Points:_ 5

```python
class KNeighborsModel:
    def __init__(self, n_neighbors: int = 5, weights: Literal["uniform", "distance"] = "uniform") -> None:
        """
        Args:
            n_neighbors (int): Количество ближайших соседей K.
            weights (Literal["uniform", "distance"]): Метод агрегации целевых переменных соседей.
                По умолчанию равен uniform.
        """
        self.n_neighbors = n_neighbors
        self.weights = weights

    @convert_pandas_to_numpy
    def fit(self, X: np.ndarray, y: np.ndarray) -> Self:
        """Обучает модель.

        Args:
            X (np.ndarray): Матрица размера (n, d), где n - количество объектов, d - количество признаков.
            y (np.ndarray): Вектор целевых переменных размера (n,).

        Returns:
            Self: Обученная модель.
        """
        assert X.shape[0] == y.shape[0], "`X` и `y` должны иметь одинаковую длину."

        ...

        return self

    @convert_pandas_to_numpy
    def predict(self, X: np.ndarray) -> np.ndarray:
        neighbors, distances = self._find_neighbors(X)
        return self._aggregate(neighbors, distances)

    def _find_neighbors(self, X: np.ndarray) -> tuple[np.ndarray, np.ndarray]:
        """Находит K-ближайших соседей.

        Args:
            X (np.ndarray): Матрица размера (m, d), где m - количество объектов, d - количество признаков.

        Returns:
            tuple[np.ndarray, np.ndarray]: Кортеж (neighbors, distances), где
                neighbors (m, k) - целевые переменные ближайших соседей,
                distances (m, k) - расстояния до ближайших соседей.
        """
        ...

    def _aggregate(self, neighbors: np.ndarray, distances: np.ndarray) -> np.ndarray:
        if self.weights == "uniform":
            return self._aggregate_uniform(neighbors)
        else:
            weights = ...
            return self._aggregate_weighted(neighbors, weights)

    def _aggregate_uniform(self, neighbors: np.ndarray) -> np.ndarray:
        raise NotImplementedError

    def _aggregate_weighted(self, neighbors: np.ndarray, weights: np.ndarray) -> np.ndarray:
        raise NotImplementedError
```

```python
grader.check("Task2")
```

### **Задание 3**. Метод K-ближайших соседей для задачи регрессии

_Points:_ 2

```python
class KNeighborsRegressor(KNeighborsModel):
    def _aggregate_uniform(self, neighbors: np.ndarray) -> np.ndarray:
        ...

    def _aggregate_weighted(self, neighbors: np.ndarray, weights: np.ndarray) -> np.ndarray:
        ...
```

```python
grader.check("Task3")
```

### **Задание 4.** Метрики качества регрессии. Root Mean Squared Error (RMSE)

Для оценки качества модели будем использовать корень из среднеквадратичной ошибки (root mean squared error, RMSE):

$$
\text{RMSE}(\mathbf{y}, \hat{\mathbf{y}}) = \sqrt{\frac{1}{m} \sum_{i=1}^{m} (y_i - \hat{y}_i) ^ 2}.
$$

_Points:_ 2

```python
def root_mean_squared_error(y_true: np.ndarray, y_pred: np.ndarray) -> np.ndarray:
    """Рассчитывает RMSE.

    Args:
        y_true (np.ndarray): Истинные значения целевой переменной.
        y_pred (np.ndarray): Предсказанные значения целевой переменной.

    Returns:
        np.ndarray: Значение метрики RMSE.
    """
    assert y_true.shape == y_pred.shape, "Размеры `y_true` и `y_pred` должны совпадать."
    ...
```

```python
grader.check("Task4")
```

### **Задание 5.** Обучение модели

Теперь обучите модель с параметром `n_neighbors=10` на числовых признаках и оцените ее качество

_Points:_ 2

```python
NUMERIC_FEATURES = ...
```

```python
reg = ...
y_pred = ...

print(f"RMSE:\t{root_mean_squared_error(y_test, y_pred):.3f}")
```

### **Задание 6.** Масштабирование признаков

Вспомните, почему так важно масштабировать признаки при использовании KNN. Затем реализуйте класс `StandardScaler` и примените его к нашим данным.

_Points:_ 2

```python
class StandardScaler:
    @convert_pandas_to_numpy
    def fit(self, X: np.ndarray) -> Self:
        ...

        return self

    @convert_pandas_to_numpy
    def transform(self, X: np.ndarray) -> np.ndarray:
        ...

    @convert_pandas_to_numpy
    def fit_transform(self, X: np.ndarray) -> np.ndarray:
        return self.fit(X).transform(X)
```

```python
...
```

```python
reg = ...
y_pred = ...

print(f"RMSE:\t{root_mean_squared_error(y_test, y_pred):.3f}")
```

```python
grader.check("Task6")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for Task0 (points: 0)
```
>>> assert data.shape == (5000, 8)
```


Tests for Task1 (points: 5)
```
>>> XA = np.array([[3.6, 4.7, 4.3], [4.9, 3.5, 3.6]])
>>> XB = np.array([[0.8, 2.7, 1.2], [2.6, 1.2, 4.5], [0.8, 4.9, 1.8]])
>>> D = cdist(XA, XB)
>>> assert D.shape == (2, 3)
>>> assert np.allclose(cdist(XA, XB), np.array([[4.631, 3.646, 3.759], [4.818, 3.375, 4.691]]), atol=0.001)
```


Tests for Task2 (points: 5)
```
>>> reg = KNeighborsModel(n_neighbors=1).fit(np.array([[0.3], [0.4], [0.0]]), np.array([1, 1, 0]))
>>> neighbors, distances = reg._find_neighbors(np.array([[0.1]]))
>>> assert neighbors.item() == 0
>>> assert np.isclose(distances.item(), 0.1)
```


Tests for Task3 (points: 2)
```
>>> neighbors = np.array([[0.4, 0.3, 0.2, 0.1, 0.1]])
>>> weights = 1 / np.array([[1.0, 1.0, 1.0, 3.0, 4.0]])
>>> assert np.isclose(KNeighborsRegressor()._aggregate_uniform(neighbors), 0.22)
>>> assert np.isclose(KNeighborsRegressor(weights='distance')._aggregate_weighted(neighbors, weights), 0.267, atol=0.001)
```


Tests for Task4 (points: 2)
```
>>> y_true = np.array([0.7, -1.5, 1.4, -1.1])
>>> y_pred = np.array([1.0, 1.7, -1.6, 0.2])
>>> assert np.isclose(root_mean_squared_error(y_true, y_pred), 2.292, atol=0.001)
>>> assert np.isclose(root_mean_squared_error(y_true, y_true), 0.0)
```


Tests for Task6 (points: 2)
```
>>> assert np.allclose(StandardScaler().fit_transform(np.array([[2.7], [2.9], [2.2], [2.0]])), np.array([[0.687], [1.236], [-0.687], [-1.236]]), atol=0.001)
```
