```python
# Initialize Otter
import otter
grader = otter.Notebook("intro_to_pytorch.ipynb")
```

# **Домашнее задание 4.** Нейронные сети

В этом задании нам предстоит написать небольшую нейронную сеть и обучить ее на датасете [Fashion-MNIST](https://www.kaggle.com/datasets/zalando-research/fashionmnist).

<p style="color: red;">
  Данное задание можно выполнять как в среде с GPU, так и без графического ускорителя — обучение модели не займет много времени.<br><br>
  Перед тем как отправить блокнот с кодом на проверку, установите <code>DEVICE = "cpu"</code> и <code>MAX_EPOCHS = 1</code>.
</p>

```python
# !pip install kaggle
# !uv pip install kaggle
```

```python
import os

import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
from tqdm import tqdm

from sklearn.metrics import classification_report

import torch
import torch.nn as nn
from torch.utils.data import Dataset, DataLoader, TensorDataset, random_split

from IPython.display import clear_output
```

```python
DEVICE = "cuda" if torch.cuda.is_available() else "cpu"
# DEVICE = "cpu"  # раскомментируйте строку перед отправкой блокнота

DEVICE
```

Давайте скачаем данные с Kaggle и посмотрим на них повнимательнее. Датасет включает обучающую и тестовую подвыборки, каждая из которых представлена отдельным CSV файлом.

Вот краткое описание данных:
- каждая строка представляет отдельное изображение в развернутом виде;
- первый столбец содержит целевую переменную;
- остальные столбцы содержат яркость соответствующих пикселей от 0 до 255 (всего 28 × 28 = 784);

В датасете 10 классов:

0.  T-shirt/top;  
1.  Trouser;  
2.  Pullover;  
3.  Dress;  
4.  Coat;  
5.  Sandal;  
6.  Shirt;  
7.  Sneaker;  
8.  Bag;  
9.  Ankle boot.

```python
if not os.path.exists("data"):
    !kaggle datasets download -p data/ --unzip zalando-research/fashionmnist

!ls data/
```

```python
train_data = pd.read_csv(os.path.join("data", "fashion-mnist_train.csv"))
print(f"{train_data.shape = }")
train_data.head()
```

```python
test_data = pd.read_csv(os.path.join("data", "fashion-mnist_test.csv"))
print(f"{test_data.shape = }")
test_data.head()
```

Отделим признаки от целевых переменных

```python
X_train, y_train = train_data.drop(columns="label").values, train_data["label"].values
X_test, y_test = test_data.drop(columns="label").values, test_data["label"].values
```

Отрисуем одно изображение

```python
print(y_train[2077])

plt.imshow(X_train[2077].reshape(28, 28), cmap="gray")
plt.show()
```

---

### **Задание 1.** Создание `Dataset` и `DataLoader`

_Points:_ 5

```python
# TODO: преобразуйте данные в `torch.Tensor`. Приведите признаки к типу данных `torch.float32`
# и отмасштабируйте их, разделив на максимальное значение
...
```

```python
# TODO: завершите класс `FashionMNISTDataset`. Полезные ссылки:
# https://docs.pytorch.org/docs/stable/data.html#torch.utils.data.Dataset
class FashionMNISTDataset(Dataset):
    def __init__(self, X: torch.Tensor, y: torch.Tensor) -> None:
        """Инициализирует новый датасет.

        Args:
            X (torch.Tensor): признаки, тензор размера (n, d), где
                n - количество объектов,
                d - количество признаков.
            y (torch.Tensor): целевые переменные, тензор размера (n,).
        """
        super().__init__()

        # TODO: сохраните данные внутри датасета
        ...

    def __getitem__(self, index: int) -> tuple[torch.Tensor, torch.Tensor]:
        """Возвращает один пример (признаки и целевую переменную) по индексу.

        Args:
            index (int): индекс примера.

        Returns:
            tuple[torch.Tensor, torch.Tensor]: кортеж вида (признаки, целевая переменная).
        """
        # TODO: верните необходимые значения
        ...

    def __len__(self) -> int:
        """Возвращает количество примеров в датасете.

        Returns:
            int: размер датасета.
        """
        # TODO: верните размер датасета
        ...
```

```python
# TODO: используйте класс `FashionMNISTDataset` для создания `train_dataset`
train_dataset = ...

# TODO: используйте класс `TensorDataset` для создания `test_dataset`. Полезные ссылки:
# https://docs.pytorch.org/docs/stable/data.html#torch.utils.data.TensorDataset
test_dataset = ...

len(train_dataset), len(test_dataset)
```

```python
# TODO: отделите 20% от обучающего датасета для валидации. Полезные ссылки:
# https://docs.pytorch.org/docs/stable/data.html#torch.utils.data.random_split
train_dataset, val_dataset = ...

len(train_dataset), len(val_dataset)
```

```python
# Создайте `DataLoader` для каждого из датасетов. Используйте `batch_size=256`
train_dataloader = ...
val_dataloader = ...
test_dataloader = ...
```

```python
batch = next(iter(train_dataloader))

X, y_true = batch
X.shape, y_true.shape
```

```python
grader.check("Task1")
```

---

### **Задание 2.** Многослойный перцептрон

Напишите двухслойный перцептрон со следующей структурой:

- линейный слой, создающий 256 скрытых признаков;
- нелинейность, например ReLU;
- линейный слой.

При желании после ReLU можно добавить слой [Dropout](`https://docs.pytorch.org/docs/stable/generated/torch.nn.Dropout.html`). Это один из способов регуляризации в нейронных сетях.

_Points:_ 5

```python
class MLPClassifier(nn.Module):
    def __init__(self, in_features: int, num_classes: int) -> None:
        super().__init__()

        ...

    def forward(self, x: torch.Tensor) -> torch.Tensor:
        ...
```

```python
grader.check("Task2")
```

---

### **Задание 3.** Цикл обучения

Наконец мы готовы приступить к обучению нашей модели.

_Points:_ 5

```python
MAX_EPOCHS = 50
# MAX_EPOCHS = 1  # раскомментируйте строку перед отправкой блокнота
```

```python
# TODO: создайте модель и перенесите ее на `DEVICE`
model = ...

# TODO: создайте оптимизатор, не забудьте передать в него параметры модели и скорость обучения
optimizer = ...

train_loss_history: list[float] = []
val_accuracy_history: list[float] = []

for epoch in range(MAX_EPOCHS):
    model.train()
    train_loss: list[float] = []

    for X, y_true in tqdm(train_dataloader, desc="Training"):
        # TODO: перенесите `X` и `y_true` на `DEVICE`
        X, y_true = ...

        # TODO: сделайте предсказание
        logits = ...

        # TODO: обнулите градиент с прошлого шага
        ...

        # TODO:  вычислите значение функции потерь. Поскольку мы решаем задачу многоклассовой
        # классификации, будем использовать cross entropy loss. Полезные ссылки:
        # https://docs.pytorch.org/docs/stable/generated/torch.nn.CrossEntropyLoss.html
        # https://docs.pytorch.org/docs/stable/generated/torch.nn.functional.cross_entropy.html
        loss = ...

        # TODO: вычислите градиент функции потерь относительно параметров модели
        ...

        # TODO: сделайте один шаг обновления весов модели
        ...

        train_loss.append(loss.item())

    train_loss: float = sum(train_loss) / len(train_loss)
    train_loss_history.append(train_loss)

    model.eval()
    val_accuracy: list[float] = []

    with torch.no_grad():
        for X, y_true in tqdm(val_dataloader, desc="Validation"):
            # TODO: перенесите `X` и `y_true` на `DEVICE`
            X, y_true = ...

            # TODO: сделайте предсказание
            logits = ...

            # TODO: определите класс, к которому модель отнесла каждый объект
            y_pred = ...

            # TODO: вычислите долю правильных ответов
            accuracy = ...

            val_accuracy.append(accuracy)

    val_accuracy: float = sum(val_accuracy) / len(val_accuracy)
    val_accuracy_history.append(val_accuracy.item())

    clear_output(wait=True)
    print(f"[{epoch:02d}/{MAX_EPOCHS}] train-loss={train_loss:.4f} | val-accuracy={val_accuracy:.4f}")
```

```python
fig, axs = plt.subplots(1, 2)
fig.set_size_inches(12, 4)

axs[0].plot(train_loss_history)
axs[0].set_xlabel("Epoch")
axs[0].set_ylabel("Train Loss")

axs[1].plot(val_accuracy_history)
axs[1].set_xlabel("Epoch")
axs[1].set_ylabel("Val Accuracy")

plt.show()
```

---

### **Задание 4.** Инференс модели

_Points:_ 5

```python
model.eval()

all_y_true: list[torch.Tensor] = []
all_y_pred: list[torch.Tensor] = []

with torch.no_grad():
    for X, y_true in tqdm(test_dataloader, desc="Inference"):
        logits = model(X).to(DEVICE)
        y_pred = logits.argmax(dim=-1)

        all_y_true.append(y_true.cpu())
        all_y_pred.append(y_pred.cpu())

all_y_true = torch.cat(all_y_true)
all_y_pred = torch.cat(all_y_pred)
```

```python
print(classification_report(all_y_true.numpy(), all_y_pred.numpy()))
```

```python
grader.check("Task4")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for Task1 (points: 5)
```
>>> assert isinstance(X_train, torch.FloatTensor)
>>> assert 0 <= X_train.min() and X_train.max() <= 1, 'признаки необходимо отмасштабировать'
>>> assert isinstance(X_test, torch.FloatTensor)
>>> assert 0 <= X_test.min() and X_test.max() <= 1, 'признаки необходимо отмасштабировать'
>>> assert isinstance(train_dataset, torch.utils.data.Subset)
>>> assert isinstance(val_dataset, torch.utils.data.Subset)
>>> assert isinstance(test_dataset, TensorDataset)
>>> assert isinstance(train_dataloader, DataLoader)
>>> assert train_dataloader.batch_size == 256
>>> assert isinstance(train_dataloader.sampler, torch.utils.data.RandomSampler), 'стоит использовать `shuffle=True`'
>>> assert isinstance(val_dataloader, DataLoader)
>>> assert val_dataloader.batch_size == 256
>>> assert isinstance(test_dataloader, DataLoader)
>>> assert test_dataloader.batch_size == 256
```


Tests for Task2 (points: 5)
```
>>> model = MLPClassifier(2, 3)
>>> x = torch.randn(1, 2)
>>> assert model(x).shape == (1, 3)
```


Tests for Task4 (points: 5)
```
>>> assert True
```
